## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

import numpy as np

a = 100
b = 0.1
c = 0.2

d0 = a*b + a*c
d1 = a*(b+c)

print(d0 == d1)
print(d1 - d0)
# This is because the binary approximation of 0.3 is difference of the sum of that of 0.1 and 0.2, 
# in python's 64bit float system
# We can try 32 bit and 16 bit to see the results
b = np.float16(0.1)
c = np.float16(0.2)
d0 = a*b + a*c
d1 = a*(b+c)
print("16 bit error:")
print(d1 - d0)

b = np.float32(0.1)
c = np.float32(0.2)
d0 = a*b + a*c
d1 = a*(b+c)
print("32 bit error:")
print(d1 - d0)
```

Output:
```
False
3.552713678800501e-15
16 bit error:
-0.01220703125
32 bit error:
7.450580596923828e-07
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1

# YOUR CODE HERE
f0 = (x+y)**6
f1 = (x**6) + 6*(x**5)*(y**1) + 15*(x**4)*(y**2) + 20*(x**3)*(y**3) + 15*(x**2)*(y**4) + 6*(x**1)*(y**5) +(y**6)

print("Compact form: ", f0)
print("Expanded form: ", f1)
print(f1-f0)
```

Output:
```
Compact form:  65944160.60120103
Expanded form:  65944160.601201
-2.9802322387695312e-08
```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

# YOUR CODE HERE
f0 = (x+y)**6
f1 = (x**6) + 6*(x**5)*(y**1) + 15*(x**4)*(y**2) + 20*(x**3)*(y**3) + 15*(x**2)*(y**4) + 6*(x**1)*(y**5) +(y**6)

print("Compact form: ", f0)
print("Expanded form: ", f1)
print(f1-f0)
```

Output:
```
Compact form:  9.999999999999788e-07
Expanded form:  9.958166629076004e-07
-4.183337092378376e-09
```

```python
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
import math
x = 1e7

original_form = 1/(math.sqrt(x**2 -1) - x)
print(original_form)

new_form = - math.sqrt(x**2 - 1) - x
print(new_form)

```

Output:
```
-19884107.85185185
-19999999.999999948
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
x = 1e9

#original_form = 1/(math.sqrt(x**2 -1) - x)
#print(original_form)

new_form = - math.sqrt(x**2 - 1) - x
print(new_form)

original_form_denominator = (math.sqrt(x**2 -1) - x)
print(original_form_denominator)

```

Output:
```
-2000000000.0
0.0
```