## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
"""
def area(x0, y0, x1, y1, x2, y2):
    return abs(x0*(y1-y2) + x1*(y2-y0) + x2*(y0-y1))/2

x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0
A = area(x0, y0, x1, y1, x2, y2)
print( round(A - 3.0, 10) == 0.0 )
"""

first_vertice = (0.0, 0.0)
second_vertice = (0.0, 2.0)
third_vertice = (3.0, 0.0)

def better_area(v0, v1, v2):
    return abs(v0[0]*(v1[1]-v2[1]) + v1[0]*(v2[1]-v0[1]) + v2[0]*(v0[1]-v1[1]))/2
A = better_area(first_vertice, second_vertice, third_vertice)
print(A)
```

Output:
```
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first_vertice = (0.0, 0.0)
second_vertice = (0.0, 3.0)
third_vertice = (4.0, 3.0)
fourth_vertice = (4.0, 0.0)
fifth_vertice = (2.0, 0.0) # kinda a dummy vertice here for the testing so I know the area

my_shape = [first_vertice, second_vertice, third_vertice, fourth_vertice, fifth_vertice]

def area(shape):
    A = 0
    for i in range(len(shape)-1):
        A += shape[i][0]*shape[i+1][1] - shape[i+1][0]*shape[i][1]
    return abs(A/2)

area(my_shape)
```

Output:
```
12.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    # YOUR CODE HERE
    z = []
    for i in range(len(x)):
        z += [x[i] + y[i]]
    return z

a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
print( c == [-2, 11.3, -20, 8] )
```

Output:
```
True
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
print( [a[i] + b[i] for i in range(len(a))] ) # the first method that came to my mind

# OR

print( [x+y for x,y in list(zip(a,b))] )
```

Output:
```
[-2, 11.3, -20, 8]
[-2, 11.3, -20, 8]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
names_to_abbrev = {'St Catharine\'s': 'CTH',
                   'Christ\'s': 'CHR', 
                   'Churchill': 'CHU', 
                   'Clare':'CL', 
                   'Clare Hall':'CLH', 
                   'Corpus Christi': 'CC'}
print(names_to_abbrev, '\n') # Testing this list

"Inversing the list"
abbrev_to_names = {}
for name, abbrev in names_to_abbrev.items():
    abbrev_to_names[abbrev] = name
print('Dictionary from college abbrev to name:\n', abbrev_to_names)
    
"Sorting the abbreviation list"
abbrev_list = [abbrev for abbrev in abbrev_to_names]
abbrev_list.sort()
print('\nSorted abbreviation list:\n', abbrev_list)
```

Output:
```
{"St Catharine's": 'CTH', "Christ's": 'CHR', 'Churchill': 'CHU', 'Clare': 'CL', 'Clare Hall': 'CLH', 'Corpus Christi': 'CC'} 

Dictionary from college abbrev to name:
 {'CTH': "St Catharine's", 'CHR': "Christ's", 'CHU': 'Churchill', 'CL': 'Clare', 'CLH': 'Clare Hall', 'CC': 'Corpus Christi'}

Sorted abbreviation list:
 ['CC', 'CHR', 'CHU', 'CL', 'CLH', 'CTH']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
cambridge_colleges = {'St Catharine\'s': ('CTH', 1473, 768),
                   'Christ\'s': ('CHR', 1505, 689), 
                   'Churchill': ('CHU', 1960, 845), 
                   'Clare': ('CL', 1326, 808), 
                   'Clare Hall': ('CLH', 1966, 249), 
                   'Corpus Christi': ('CL', 1352, 553)}

"Find college with the greatest number of students and print its abbreviation"
greatest_num_students = 0
for n, l in cambridge_colleges.items():
    if l[2] > greatest_num_students:
        greatest_num_students_college_abbrev = l[0]
        greatest_num_students = l[2]
print('The college with the greatest number of students in this list is: ', 
      greatest_num_students_college_abbrev)



"Find oldest college and print its number of students and its abbreviation"
oldest_college_year = 2020
for n, l in cambridge_colleges.items():
    if l[1] < oldest_college_year:
        oldest_college_abbrev = l[0]
        oldest_college_numstudents = l[2]
        oldest_college_year = l[1]
print('Oldest College in the list is ', oldest_college_abbrev, 
      ', with number of students: ', oldest_college_numstudents)

```

Output:
```
The college with the greatest number of students in this list is:  CHU
Oldest College in the list is  CL , with number of students:  808
```